```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

```python
# Import me!
import numpy as np
import matplotlib.pyplot as plt
```

---
# **Homework 2: Getting started with Python (34 pts)**

### Due Friday, October 16th at 11:59 pm

### Point Breakdown
| **Question #** | **Points** | **Manual Grading** | 
| -------------- | ---------- | ---------- |
| 1a             | 3         | No        |
| 1b             | 3          | No         |
| 1c             | 4          | No        |
| 2a            | 3          | No         |
| 2b            | 3          | No        |
| 2c             | 1          | No         |
| 2d             | 4          | No         |
| 2e             | 4          | No         |
| 2f             | 4          | No         |
| 3             | 5          | Yes         |
#### Total points: 16

---
## Question 1: Defining Variables (10 pts)

There are an estimated 20,000,000,000,000,000 (twenty quadrillion) ants on Earth and slightly over 8,000,000,000 (eight billion) humans. Let's say the average mass of an ant is 0.005 g and the average mass of a human is 62 kg.

### Question 1a (3 pts)

Define four variables: `ant_pop`, `hum_pop`, `ant_mass_g`, and `hum_mass_kg`, assigning them to the appropriate numeric values.

**Note:** *In order for Python to properly read a number, it cannot include commas.* 

```python
ant_pop = ...
hum_pop = ...
ant_mass_g = ...
hum_mass_kg = ...
```

```python
grader.check("q1a")
```

### Question 1b (3 pts)

Using the variables you defined, write an expression that computes the number of ants per human, assigning it to the variable `ants_per_human`.

```python
ants_per_human = ...
print("For every human, there are", ants_per_human, "ants.")
```

```python
grader.check("q1b")
```

### Question 1c (4 pts)

Write another expression for the overall ant:human mass ratio on Earth. Start by defining a new variable `hum_mass_g` which converts `hum_mass_kg` to grams. Then define `ants_to_hums_mass` as the ratio of the total mass of all ants to the total mass of all humans.

```python
hum_mass_g = ...
ants_to_hums_mass = ...
ants_to_hums_mass
```

You should have found that the total mass of all ants is about 1/5 the total mass of all humans. Alert your local government if the ratio approaches 1. 

**Note:** *When defining variables, be sure to use readable, intelligible names so that you can understand their purpose long after you've defined them. Try to keep the names short if possible, especially if they'll be used often.*

```python
grader.check("q1c")
```

---
## Question 2: Lists and Slicing (22 pts)

### Question 2a (3 pts)
Using the list `bank` (defined below), add the correct indexes to the expression:
```
"A " + bank[][] + " flower has " + bank[][] + " " + bank[][] + " and " + bank[][] + " " + bank[][]
```
So that the concatenated string produced reads as:
```
"A Brassica flower has 4 petals and 6 stamens"
```
**Note:** *Copy/paste the empty expression from the raw cell below into the answer cell. Fill in the brackets with the appropriate integers. Make sure not to add any extra spaces to the expression.*


"A " + bank[][] + " flower has " + bank[][] + " " + bank[][] + " and " + bank[][] + " " + bank[][]

```python
# run this cell to assign bank
# please don't alter the list
bank = [['Melica', 'Brassica', 'Calochortus'], ['4', '1', '2', '1', '6'], ['stem', 'stamens', 'pistil', 'sepals', 'petals']]
```

```python
# construct your sentence
sentence = ...
print(sentence)
```

```python
grader.check("q2a")
```

### Question 2b (3 pts)
Split the word *bioinformatician* into prefix, root, and suffix using slicing. The three parts should be `bio`, `informatic`, and `ian`. <br>
For the root and suffix, use a negative indicies as part of your solution.<br>


```python
word = 'bioinformatician'
prefix = ...
root = ...
suffix = ...
print(prefix, "-", root, "-", suffix)
```

```python
grader.check("q2b")
```

### Question 2c (1 pt)
Make the output of `word` from question 3c plural by concatenating it with a simple string that is the letter `s`, giving you `bioinformaticians` (you can use the `+` operator for this).

```python
plural = ...
plural
```

```python
grader.check("q2c")
```

### Question 2d (4 pts)
Someone with inconsistent capitalization has recorded the names of their favorite birds in the Python list below. Use slicing to access the specified substrings. Use the `+` operator to concatenate these substrings into a secret message and assign it to the variable `secret_message`.

Get these substrings:
- the last two characters of the first bird
- the last three characters of the fifth bird
- the third and fourth characters of the sixth bird
- the second character of the third bird
- and the the tenth through twelveth characters of the fourth bird

*Hint* concatenate substrings with the `+` operator like so:

- ```birdlist[0] + birdlist[1] + ...```


```python
# run this cell to assign birdlist
# please don't alter the list
birdlist = ['QUAIL', 'Blue JAY', 'grebe', 'Red-rumped Swallow', 'DovE', 'rObin']
```

```python
# decode the message here
secret_message = ...
secret_message
```

```python
grader.check("q2d")
```

### Question 2e (4 pts)

Given the list `cities`, assign the following variables using list indexing.

`first_five_cities` - the first five cities in the list.<br>
`last_three_cities` - the last three cities in the list.<br>
`reversed_list_of_cities` - the list in reverse order.<br>
`every_other_city` - every other city in the list, starting with the first (New York).<br>
`cities_excluding_first_and_last` - a list that excludes the first and last city.<br>

```python
# run this cell to assign cities
# please don't alter the list
cities = ["New York", "Los Angeles", "Chicago", "Houston", "Phoenix", "Philadelphia", "San Antonio", "San Diego", "Dallas", "San Jose"]
```

```python
# assign new variables
first_five_cities = ...
last_three_cities = ...
reversed_list_of_cities = ...
every_other_city = ...
cities_excluding_first_and_last = ...

# print results
print(f"First five cities: {first_five_cities}\n")
print(f"Last three cities: {last_three_cities}\n")
print(f"Reversed: {reversed_list_of_cities}\n")
print(f"Every other city: {every_other_city}\n")
print(f"Excluding first and last: {cities_excluding_first_and_last}\n")
```

```python
grader.check("q2e")
```

### Question 2f (4 pts)
You study environmental factors that affect lobster growth and have been given the length (in cm) of collected crustaceans in a Python list.
For each variable below, write a formula that produces the desired information.<br><br>
`num_lobsters` - the total number of lobsters in the list.<br>
`max_length` - the length of the longest lobster.<br>
`mean_length` - average lobster length.<br>
`subset` - a list of every other lobster from the 7th to the 77th.<br>
`record` - a list with one additional lobster. A record lobster was found in your study area. Add a value of 99 to the list `lobsters`.

```python
# run this cell to assign lobsters
# please don't alter the list
lobsters = [29, 44, 55, 56, 52, 55, 33, 43, 62, 68, 52, 30, 19, 56, 28, 69, 38, 70, 36, 30, 52, 28, 44, 40, 30, 22, 49, 24, 20, 55, 27, 53, 26, 33, 61, 56, 49, 23, 15, 60, 54, 55, 40, 54, 57, 26, 68, 53, 24, 36, 46, 36, 67, 68, 53, 29, 30, 55, 70, 33, 55, 28, 26, 37, 27, 62, 38, 37, 43, 51, 47, 55, 39, 33, 17, 58, 57, 27, 53, 49, 29, 48, 70, 46, 28, 40, 62, 41, 59, 17, 40, 44, 44, 25, 22, 51, 61, 24, 38, 55, 56, 28, 67, 27, 44, 31, 36, 38, 24, 43, 63, 69, 67, 18, 29, 33, 42, 63, 28, 47, 42, 24, 38, 44, 60, 16, 40, 57, 28, 35, 54, 58, 38, 62, 59, 36, 29, 17, 66, 55, 46, 66, 36, 25, 32, 56, 16, 36, 59, 23, 46, 53, 32, 23, 43, 22, 55, 30, 50, 28, 57, 52, 69, 53, 33, 51, 50, 52, 43, 35, 62]
```

```python
# assign new variables
num_lobsters = ...
max_length = ...
mean_length = ...
subset = ...
record = ...

# print results
print("Number of lobsters studied:", num_lobsters)
print("Longest lobster:", max_length, "cm")
print("Average lobster length:", mean_length, "cm")
print("Subset =", subset)
print("Record added =", record)
```

```python
grader.check("q2f")
```

<!-- BEGIN QUESTION -->

---
## Question 3: Plotting Python Lists (5 pts)

*This question will be manually graded*

An important part of data analysis is data visualization. We can plot data stored in Python lists with the `matplotlib.pyplot` library (imported at the top of this notebook as `plt`). 

- More information about the `matplotlib` library: https://matplotlib.org/stable/users/explain/quick_start.html

To get started, we will plot some rainfall data for Corvallis. The code cell below creates two lists:

- `month` is a list of month name abbreviations in order.
- `corv_avg_monthly_rainfall` is a list of twelve floating point numbers. Each number represents the average rainfall received by Corvallis for a given month (in mm). Each value was computed using data collected between 2005 and 2015.

```python
month = ['Jan', 'Feb', 'Mar', 'Apr', 'May', 'Jun', 'Jul', 'Aug', 'Sep', 'Oct', 'Nov', 'Dec']
corv_avg_monthly_rainfall = [146.05, 96.91, 135.93, 70.54, 64.38, 32.47, 7.58, 9.35, 42.67, 79.29, 171.51, 183.29]
```

Now let's make a plot of average monthly rainfall for Corvallis. Make sure the visualization package `matplotlib.pyplot` is imported (run the import cell at the top of this notebook if you haven't already).

- The X axis of our plot should represent each of the twelve months of the year in order.
- The Y axis should represent the average monthly rainfall corresponding to each month.

Make a line plot using the following method of `matplotlib` (which we have abbreviated to `plt`) 

- ```plt.plot(x, y)```

*Hint*: Replace `x` and `y` with the appropriate lists (created above).

It's important to be able to plot multiple data layers on top of each other. By default, `matplotlib` will construct a line plot. You can specify a point plot by providing the string `'o'` as the third argument of `plt.plot()`. Construct a combined line-point plot by executing `plt.plot()` twice like so:

- ```plt.plot(x, y)```
- ```plt.plot(x, y, 'o')```

*Hint*: What happens if you switch the order of these lines?

Include labels for the X and Y axes and title the plot using the following methods of `matplotlib`.

- ```plt.xlabel('this is my x axis label')```
- ```plt.ylabel('this is my y axis label')```
- ```plt.title('this is my title')```

Inspect the plot and do some interpretation. Does the pattern you see make sense?
    

```python
# plotting
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Save your notebook and then run the last cell. Download the generated zip file and upload it to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 3)
```
>>> assert ant_pop == 20000000000000000
```
failure_message: Ant population is incorrect.
```
>>> assert hum_pop == 8000000000
```
failure_message: Human population is incorrect.
```
>>> assert ant_mass_g == 0.005
```
failure_message: Ant mass is incorrect.
```
>>> assert hum_mass_kg == 62
```
failure_message: Human mass is incorrect.


Tests for q1b (points: 3)
```
>>> assert ants_per_human == 2500000.0
```
failure_message: Use the variables ant_pop and hum_pop to divide the total ant population by the total human population.


Tests for q1c (points: 4)
```
>>> assert hum_mass_g == 62 * 1000
```
failure_message: Make sure you are converting kilograms to grams correctly.
```
>>> assert ants_to_hums_mass != 4.96
```
failure_message: The ratio should be ant:human not human:ant.
```
>>> assert ants_to_hums_mass == 0.20161290322580644
```
failure_message: To compute the total mass for a species, multiply the population and average mass. Use division to find the ratio of these total masses.


Tests for q2a (points: 3)
```
>>> assert sentence == 'A Brassica flower has 4 petals and 6 stamens'
```
failure_message: Remember Python uses zero-based indexing (meaning the 1st value is at index 0).


Tests for q2b (points: 3)
```
>>> assert prefix == word[0:3]
```
failure_message: prefix is incorrect.
```
>>> assert root == word[3:13]
```
failure_message: root is incorrect.
```
>>> assert suffix == word[-3:16]
```
failure_message: suffix is incorrect.


Tests for q2c (points: 1)
```
>>> assert plural == word + 's'
```
failure_message: Use the "+" operator to concatenate strings (there are other valid options as well).


Tests for q2d (points: 4)
```
>>> assert secret_message == 'ILovEbird S'
```
failure_message: secret message incorrect - The messag should be readable but does not follow capitalization conventions


Tests for q2e (points: 4)
```
>>> assert first_five_cities == cities[:5]
```
failure_message: first_five_cities is incorrect.
```
>>> assert last_three_cities == cities[-3:]
```
failure_message: last_three_cities is incorrect.
```
>>> assert reversed_list_of_cities == cities[::-1]
```
failure_message: reversed_list_of_cities is incorrect.
```
>>> assert every_other_city == cities[::2]
```
failure_message: every_other_city is incorrect.
```
>>> assert cities_excluding_first_and_last == cities[1:-1]
```
failure_message: cities_excluding_first_and_last is incorrect.


Tests for q2f (points: 4)
```
>>> assert num_lobsters == len(lobsters)
```
failure_message: num_lobsters incorrect.
```
>>> assert max_length == max(lobsters)
```
failure_message: max_length incorrect.
```
>>> assert mean_length == sum(lobsters) / len(lobsters)
```
failure_message: mean_length incorrect - use len() or num_lobsters and sum().
```
>>> assert subset == lobsters[6:77:2]
```
failure_message: subset incorrect - are you indexing EVERY OTHER element from 7 to 77 (meaning your selection is alternating)?
```
>>> assert record == lobsters + [99]
```
failure_message: record incorrect - record-breaking measurement was not added (or lobsters was altered).
