```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

```python
# Import me!
import os
```

---
# **Homework 1: UNIX Shell Commands (34 pts)**

### Due Friday, October 9th at 11:59 pm

In this homework, you will be working within the UNIX shell (hereafter just called "the shell"). A functional understanding of the shell will enable you to effectively manipulate your filesystem. However, the shell can be frustrating. It follows different rules than Python and is generally more difficult to read. The purpose of this notebook is to practice the most essential (and useful) shell commands covered in lecture. Don't be afraid to ask fellow classmates and course assistants for help!

**Point Breakdown:**

*All questions are automatically graded*
- Q1: 5 points
- Q2: 5 points
- Q3a: 6 points
- Q3b: 6 points
- Q4a: 6 points
- Q4b: 6 points

**Helpful Resources**
- [Shell/terminal help](https://swcarpentry.github.io/shell-novice/) :  Lessons on shell commands (ls, cd, rm, cp, mv etc.)  
- [Jupyter Lab terminal](https://jupyterlab.readthedocs.io/en/stable/user/terminal.html) : How to use the terminal within Jupyter Lab


---
## Using the UNIX Shell

In the Jupyter Hub environment, shell commands can be executed from an interactive terminal. To open a new terminal, click the `+` next to your tabs (top left), then select `Terminal`. Drag the terminal tab to one side of your screen so that this notebook and the terminal are side by side. Try to execute some commands in the terminal. Try `pwd` and `ls`.

Shell commands can also be run from this notebook. In order for a code cell to be interpreted as a series of shell commands (rather than lines of Python code), `%%bash` must be at the top. Run the cell below and inspect the output.

```python
%%bash
pwd
ls
```

Cool! We see that the default working directory is the directory where this notebook is located. But be careful. Your working directory will be reset to the default each time a new `%%bash` cell is executed. This can be confusing (see below). You should treat each `%%bash` cell as a self-contained journey that starts from the default directory each time.

```python
%%bash
pwd # starting where the notebook is located
cd .. # moving up one level
pwd # seemed to work
```

```python
%%bash
pwd # wait? I am back here now?
```

As you work through this notebook, you can execute shell commands in `%%bash` cells, but don't be afraid to execute commands in the terminal as well! Keep in mind that both methods are manipulating **the same filesystem**. So changes made using one method will be visible using the other.

---
### Relative and Absolute Paths

**Absolute Paths** are paths specified from the root of the filesystem
- ex. `/Users/billy/myproject/data/mydata.txt` or `~/myproject/data/mydata.txt`

Absolute paths are useful because they will be valid no matter where you are on your filesystem. But when explicitly typed out, absolute paths will fail on other filesystems, simply because other filesystems won't have the same directory structure.
For example, when Billy shares `myproject/` with his friend Tommy, Tommy will encounter an error telling him that `/Users/billy/` is not a valid directory.

**Relative Paths** are paths specified from your current directory
- ex. `data/mydata.txt` *assuming you are in* `/Users/billy/myproject/`

Relative paths avoid this issue (assuming working directories are standardized across filesystems). For example, even if Tommy puts `myproject/` in `/Users/tommy/Downloads/`, the path above will still be valid if executed from `myproject/`. However, relative paths fail if working directory structure is not standardized across filesystems or commands are executed from outside of the working directory.

Whether you should be using absolute or relative paths is context-dependent. **<span style="color:red">For this notebook, use either relative paths, or dynamic absolute paths (see below).</span>** Explicitly typed out absolute paths will fail when we attempt to run your submitted notebooks for grading.



```python
%%bash
# this cell sets up a dynamic symlink
# that will enable you to use absolute paths
# without causing issues on other filesystems
# don't alter this cell please

# remove old symlink (if one exists)
if [ -L ~/symhw01 ]; then
    echo "removing old symlink"
    rm ~/symhw01;
fi

# create symlink in the home directory called `symhw01` that
# points to the current directory (accessed with built in variable $PWD)
HW01_PATH=$PWD
ln -s $HW01_PATH ~/symhw01
```

```python
%%bash
# now we have flexibility when specifying paths
# don't alter this cell please

cat ~/symhw01/test.txt # we can use our handy symlink
echo ' (using symlink)'
cat test.txt # we can use relative paths
echo ' (using relative path)'
```

Relative paths may be easier for this assignment. If you choose instead to use the symlink (dynamic absolute path), use `~/symhw01` in place of the absolute path to the current directory you are working in. For example, if you want to access the file located at `/home/yourname/class/data/mydata.txt`, and your current directory is `/home/yourname/class/`, you would type `~/symhw01/data/mydata.txt`.

---
### Question 1: Basic Shell Commands (5 pts)
In a hypothetical filesystem, what is the output of the final `ls` command in the sequence shown below?

*Lines beginning with `$` are commands. Lines without `$` show the output of the previous command (if any output was produced).*

$ pwd
/Users/jamie/data
$ ls
proteins.dat
$ mkdir recombine
$ mv proteins.dat recombine/
$ cp recombine/proteins.dat ../proteins-saved.dat
$ ls

Type your solution as a python string and assign it to the variable `answer_1` below.

```python
answer_1 = ...
```

```python
grader.check("q1")
```

---
### Question 2: Pattern Matching (5 pts)
Given the list of files below, type the command that would list only the files ending in `.bam`. Assume you are already in the directory where these files are located. Try to make your final solution as general as possible.

*Hint*: Use `*` to match any character/characters. For example the command `ls *31*` would list the first four files below.

```BC-Inf-8-31-22-A.depth```

```BC-Inf-8-31-22-A.sorted.bam```

```BC-Inf-8-31-22-A.sorted.bam.bai```

```BC-Inf-8-31-22-A.sorted.sam```

```BD-Inf-9-5-22-A.depth```

```BD-Inf-9-5-22-A.sorted.bam```

```BD-Inf-9-5-22-A.sorted.bam.bai```

*Hint*: Make sure you are only listing files that *end* in `.bam`, not just ones that contain `.bam` anywhere in the filename.

To test your solution, use the `%%bash` cell below to unpack the directory `pattern-matching-practice/` (containing these files) into your current directory. Then use the second `%%bash` cell to execute your command. Does it return what you expect?

```python
%%capture
%%bash
# unpacks pattern-matching-practice
# don't alter this cell please
rm -rf pattern-matching-practice/
tar -xzvf pattern-matching-practice.tar.gz
```

```python
%%bash
# change to the directory where the files are located
cd pattern-matching-practice/
# test out your command here

```

When you have your solution, type it as a Python string and assign it to the variable `answer_2` below. For example `answer_2` = `'your command'`.

```python
answer_2 = ...
```

```python
grader.check("q2")
```

---
## Directory and File Manipulation

Questions 3 and 4 will ask you to manipulate directory structures stored in the file `hw01-data.tar.gz`. A dedicated code cell will unpack this file into your current directory (as one did in question 2). While you can use either the terminal or the existing `%%bash` cells to work out which commands you need to run, **<span style="color:red">you must enter your final sequence of shell commands into the `%%bash` cells so that when executed in order they pass all tests</span>**. To test if your commands do what they are supposed to do, execute the unpacking cell to "refresh" the directory structure, then run all `%%bash` cells in order.

---
### Question 3: Copying and Renaming (12 pts)

The cell below will unpack `hw01-data.tar.gz` into your current directory. Running the cell again will reset `hw01-data/` to its initial state. Make sure not to delete `hw01-data.tar.gz`. **<span style="color:red">If you accidentally mess up your directory structure or want to start from scratch, re-run this cell!</span>**

```python
%%capture
%%bash
# unpacks hw01-data/
# don't alter this cell please
rm -rf hw01-data/
tar -xzvf hw01-data.tar.gz
```

After running the cell above, you should see the directory `hw01-data/` in your current directory. Use the `%%bash` cell below to explore it. *Hint*: Use the `ls` or `tree` commands.

```python
%%bash
# the notebook and the tree command disagree about how output should be encoded
# this command fixes that
export LC_ALL=C
# explore hw01-data/

```

#### Question 3a (6 pts)
Let's back up some of our data. Make a directory `backups/` inside `hw01-data/`. Within `hw01-data/datasets/animal-data/`, there is a directory called `butterfly/`. Copy the `butterfly/` directory (and all its contents) to `hw01-data/backups/`.

```python
%%bash
# enter your commands here
...
```

```python
grader.check("q3a")
```

#### Question 3b (6 pts)
Let's also back up these notes. Make a directory `other/` inside `hw01-data/backups/`. Then copy `notes.txt` from `hw01-data/datasets/` to `hw01-data/backups/other/`. Rename the backed-up copy `notes-backup.txt`.

```python
%%bash
# enter your commands here
...
```

```python
grader.check("q3b")
```

---
### Question 4: Moving and Removing (12 pts)

Remember, **<span style="color:red">If you accidentally mess up your directory structure, re-run the unpacking cell at the top of Question 3a to reset it!</span>** (To get back to where you were, run the subsequent `%%bash` cells in order).

#### Question 4a (6 pts)
Let's organize some files. In the `hw01-data/lakers/` directory there are eight `.csv` files which contain play-by-play shooting data from the first two games of the Los Angeles Lakers 2008-2009 season. They played the Portland Trail Blazers (POR) game one and Los Angeles Clippers (LAC) game two (winning both).

Make a new directory called `portland/` inside `hw01-data/datasets/`. Move all the `.csv` files beginning with `POR` into the new `portland/` directory. Try to move all files with a single command.

*Hint*: Revisit Question 2 for pattern matching practice.

```python
%%bash
# enter your commands here
...
```

```python
grader.check("q4a")
```

#### Question 4b (6 pts)
Let's practice removing files. First, copy the `lakers/` directory and all of its contents into `hw01-data/backups/`. Then remove `LAC_period_4.csv` from the original `lakers/` directory. Finally, remove the original `butterfly/` directory and all of its contents from `hw01-data/datasets/animal-data/`.

```python
%%bash
# enter your commands here
...
```

```python
grader.check("q4b")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Save your notebook and then run the last cell. Download the generated zip file and upload it to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 5)
```
>>> assert 'answer_1' in vars().keys()
```
failure_message: It looks like you have deleted the 'answer_1' variable. Assign your answer as a string to a python variable named 'answer_1'.
```
>>> assert answer_1 != ...
```
failure_message: It looks like you didn't assign anything to the variable 'answer_1'.
```
>>> assert isinstance(answer_1, str)
```
failure_message: It looks like your answer is not a string. To make it a string put quotes around it (e.g. answer_1 = 'answer').
```
>>> assert answer_1 == 'recombine' or answer_1 == 'recombine/'
```
failure_message: You did not enter the correct file and/or directory name(s).


Tests for q2 (points: 5)
```
>>> assert 'answer_2' in vars().keys()
```
failure_message: It looks like you have deleted the 'answer_2' variable. Assign your answer as a string to a python variable named 'answer_2'.
```
>>> assert answer_2 != ...
```
failure_message: It looks like you didn't assign anything to the variable 'answer_2'.
```
>>> assert isinstance(answer_2, str)
```
failure_message: It looks like your answer is not a string. To make it a string put quotes around it (e.g. answer_2 = 'answer').
```
>>> assert 'ls' in answer_2
```
failure_message: It looks like the wrong command is being used. Check the lecture notes or the shell documentation linked earlier in this notebook to find the correct command to use.
```
>>> assert '*' in answer_2
```
failure_message: It looks like your command does not capture multiple files. Refer to the hint in this question.
```
>>> assert 'pattern-matching-practice' in os.listdir()
```
failure_message: There is no pattern-matching-practice/ directory in your current directory - Make sure you run the %%bash cell to unpack it.
```
>>> answer_2_result = sorted(os.popen('cd pattern-matching-practice && ' + answer_2).read().rstrip().split('\n'))
>>> assert answer_2_result == ['BC-Inf-8-31-22-A.sorted.bam', 'BD-Inf-9-5-22-A.sorted.bam']
```
failure_message: It looks like your command does not capture multiple files or the correct files
```
>>> assert answer_2 == 'ls *.bam'
```
failure_message: It looks like your command captures the correct files. But try and come up with a general command that will always capture all .bam files.


Tests for q3a (points: 6)
```
>>> assert 'hw01-data' in os.listdir()
```
failure_message: There is no hw01-data/ directory in your current directory - Make sure you run the %%bash cell to unpack it.
```
>>> assert 'backups' in os.listdir('hw01-data/')
```
failure_message: There is no backups/ directory in hw01-data/
```
>>> assert 'butterfly' in os.listdir('hw01-data/backups/')
```
failure_message: There is no butterfly/ directory in hw01-data/backups/
```
>>> assert 'butterfly' in os.listdir('hw01-data/datasets/animal-data/')
```
failure_message: There is no butterfly/ directory in hw01-data/datasets/animal-data/ - Did you delete the original data?
```
>>> butterfly_data = ['butterfly00.fastq', 'butterfly01.fastq', 'butterfly02.fastq', 'butterfly03.fastq', 'butterfly04.fastq']
>>> assert sorted(os.listdir('hw01-data/datasets/animal-data/butterfly/')) == butterfly_data
```
failure_message: The original data has been altered or deleted
```
>>> assert os.listdir('hw01-data/backups/butterfly/') == os.listdir('hw01-data/datasets/animal-data/butterfly/')
```
failure_message: The copied data does not match the original data


Tests for q3b (points: 6)
```
>>> assert 'other' in os.listdir('hw01-data/backups/')
```
failure_message: There is no other/ directory in hw01-data/backups/
```
>>> assert 'notes-backup.txt' in os.listdir('hw01-data/backups/other/')
```
failure_message: There is no notes-backup.txt file in hw01-data/backups/other/ - Did you rename the copied file?
```
>>> assert 'notes.txt' in os.listdir('hw01-data/datasets/')
```
failure_message: There is no notes.txt file in hw01-data/datasets/  - Did you delete the original file?
```
>>> notes1 = open('hw01-data/datasets/notes.txt').readlines()
>>> notes2 = open('hw01-data/backups/other/notes-backup.txt').readlines()
>>> assert notes1 == notes2
```
failure_message: The contents of notes-backup.txt does not match the original contents of notes.txt


Tests for q4a (points: 6)
```
>>> assert 'portland' in os.listdir('hw01-data/datasets/')
```
failure_message: There is no portland/ directory in hw01-data/datasets/
```
>>> POR_files = ['POR_period_1.csv', 'POR_period_2.csv', 'POR_period_3.csv', 'POR_period_4.csv']
>>> assert POR_files == sorted(os.listdir('hw01-data/datasets/portland/'))
```
failure_message: Contents of hw01-data/datasets/portland/ are not correct - Make sure you moved only 4 files beginning with POR
```
>>> LAC_files = ['LAC_period_1.csv', 'LAC_period_2.csv', 'LAC_period_3.csv', 'LAC_period_4.csv']
>>> assert LAC_files == sorted(os.listdir('hw01-data/lakers/'))
```
failure_message: Contents of hw01-data/lakers/ are incorrect - Did you move all of the files? Did you copy the files instead of moving them?


Tests for q4b (points: 6)
```
>>> assert 'lakers' in os.listdir('hw01-data/backups/')
```
failure_message: There is no lakers/ directory in hw01-data/backups/
```
>>> LAC_files = ['LAC_period_1.csv', 'LAC_period_2.csv', 'LAC_period_3.csv', 'LAC_period_4.csv']
>>> assert LAC_files == sorted(os.listdir('hw01-data/backups/lakers/'))
```
failure_message: The contents of hw01-data/backups/lakers/ don't match the original
```
>>> assert 'LAC_period_4.csv' not in os.listdir('hw01-data/lakers/')
```
failure_message: LAC_period_4.csv hasn't been deleted from hw01-data/lakers/
```
>>> LAC_files_reduced = ['LAC_period_1.csv', 'LAC_period_2.csv', 'LAC_period_3.csv']
>>> assert LAC_files_reduced == sorted(os.listdir('hw01-data/lakers/'))
```
failure_message: Other files are missing from hw01-data/lakers/ - Did you remove more than just LAC_period_4.csv?
```
>>> assert 'butterfly' in os.listdir('hw01-data/backups/')
```
failure_message: butterfly/ is missing from hw01-data/backups/ - Make sure to remove only the original directory in hw01-data/datasets/animal-data/
```
>>> assert 'butterfly' not in os.listdir('hw01-data/datasets/animal-data/')
```
failure_message: butterfly/ has not been removed from hw01-data/datasets/animal-data/
