```python
# Initialize Otter
import otter
grader = otter.Notebook("hw0_master.ipynb")
```

```python
import numpy as np
```

# **Homework 0: Working on the Jupyter Notebook and Python Basics**
## Due Friday, October 2nd at 11:59pm
This problem set is meant an introduction to Python using Jupyter Notebooks and to the mechanics of submitting assignments in BDS 310.<br>
*You will get full credit for this problem set if you attempt to answer questions and submit it via Gradescope.*<br>

Remember - to submit, you need to download a .zip file to your personal computer and then submit via Gradescope tab on Canvas or via Gradescope web site.<br>


Overview:
1. [Jupyter Notebooks](#jup_note)
    - 1.1 [Text Cells](#txt_cell)
    - 1.2 [Code Cells](#code_cells)
    - 1.3 [Writing Jupyter Notebooks](#write)
    - 1.4 [Errors](#errors)
    - 1.5 [The Kernel](#kernel)
2. [Python Basics](#python_basics)
    - 2.1 [Print & Comment](#print_comment)
    - 2.2 [Data Types](#data_types)
        - 2.2.1 [Strings](#strings)
        - 2.2.2 [Integers & Floating Point Numbers](#int_float)
    - 2.3 [Expressions](#exp)
    - 2.4 [Variable Assignment](#var)
3. [Question 3](#q3)
4. [Question 4](#q4)

<a id="jup_note"></a>
# 1. Jupyter Notebooks
This webpage is called a Jupyter Notebook. A notebook is a place to write programs and view their results, and also to write text.

<a id="txt_cell"></a>
## 1.1. Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings. You won't need to learn Markdown for this class, but you might want to.

After you edit a text cell, click the **run cell** button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

**Question 1.1**<br> This paragraph is in its own text cell.  Try editing it so that this sentence is the last sentence in the paragraph, and then click the "run cell" ▶| button or hold down `shift` + `return`.  This sentence, for example, should be deleted.  So should this one.

<a id='code_cells'></a>
## 1.2. Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it. It'll be highlighted with a green or blue rectangle. Next, either press ▶| or hold down `shift` + `return`.

Try running this cell:

```python
print("Hello, World!")
```

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

The fundamental building block of Python code is an expression. Cells can contain multiple lines with various expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

**Question 1.2** Change the cell above so that it prints out:

    First this line,
    then the whole world,
    and then this one.

*BONUS* If you can, try to implement the Earth symbol into the print statement. If you get stuck, chat with one of your classmates

<a id='write'></a>
## 1.3. Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a code cell. You can change it to a text cell by clicking inside it so it's highlighted, clicking the drop-down box next to the `restart` button in the menu bar, and choosing "Markdown".

**Question 1.3** <br>Add a code cell below this one. Write the following line into the code cell:
   
    print("A whole new cell!")

Run your cell to verify that it works.

```python

```

<a id='errors'></a>
## 1.4. Errors
Python is a language, and like natural human languages, it has rules. It differs from natural language in two important ways:
1. The rules are *simple*. You can learn most of them in a few weeks and gain reasonable proficiency with the language in a quarter.
2. The rules are *rigid*. If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes. A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes. When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong. Errors are OK. You just have to find the source of the problem, fix it, and move on.

We have made an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

```python

```

**Note:** In the toolbar, there is the option to click `Cell > Run All`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="error.jpg"/>

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  "`EOF`" means "end of file," so the message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all to program effectively. If you see a cryptic message like this, you can often get by without deciphering it. (Of course, if you're frustrated, ask a neighbor or a staff member for help.)

Try to fix the code above so that you can run the cell and see the intended message instead of an error.

You will likely encounter various Python errors and exceptions. Below are some resources to understand these errors.

https://www.geeksforgeeks.org/errors-and-exceptions-in-python/<br>
https://docs.python.org/3/tutorial/errors.html

<a id='kernel'></a>
## 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `In [...]`. Before you run the cell, you'll see `In [ ]`. When the cell is running, you'll see `In [*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `In [1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is slow and unresponsive, or your kernel loses its connection.If this happens, try the following steps:
1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.

```python

```

## 1.6 Built-In Tests
Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. Please don't change the contents of the test cells. 

Running the cell directly after it will test whether you have assigned the correct data or expression. If you haven't, this test will tell you the correct answer. (Sometimes the tests will give hints about what went wrong...)

Let's try an example question for what this might look like in a typical homework assignment...

**Question 1.6**: <br>

What is 2+2?

Example Question for practicing using the `grader.check()` to help answer this question

- Part 1: Test a Wrong Answer:<br><br>
In the above question, it asks you what 2+2 is. We know the answer, but let us start by inputting the value of 5 to the right of the `=` for the variable test_answer.<br><br>
After assigning this value, click the "run cell" ▶| button or hold down `shift` + `return`.<br><br>
Next, after running  your answer cell, you can execute the cell containing `grader.check()`. This will run all of the tests for the question and inform you of errors or if you are correct.

- Part 2: Test a Right Answer:<br><br>
Perform the same procedure as before, but change the value of 5 to be the correct answer for 2+2. <br><br>
After running the `grader.check()` cell you should receive an indication that you got the question correct and can move on.


```python
test_answer = ...
```

```python
grader.check("q1.6")
```

<a id="python_basics"></a>
# 2. Python Basics

<a id='print_comment'></a>
## 2.1 Print and Comment 
As we saw before, the `print` function is a method in Python to aid in debugging your code.

In addition to print statements for continuous testing, comments can concisely explain your thought process. It is important to include any relevant units related to constants or computed values within these comments. Python comments are displayed with  `#` and when the program is executed they will not appear in the output.

Below is an example of one such print statement with a comment. For the following exercise,  continue to experiment adding print statements and comments.

```python
# This is a comment
print("This is not commented")
print(2024)
print("...hint?")
# print("This is commented")
# If you comment a print statement it will not appear in the output
```

In addition to the `print` function the last value in the cell block will also print in the cell output automatically. See below:

```python
print(2)
3
4
```

Above, you should see that 4 is the value of the last expression, 2 is printed, but 3 is lost forever because it was neither printed nor last.

You don't want to print everything all the time anyway.  But if you feel sorry for 3, change the cell above to print it.

<a id='data_types'></a>
# 2.2 Data Types:

Now we will briefly cover the python data types that will be used throughout this course. 

<a id="strings"></a>
### 2.2.1 Strings
*Strings* are characters in a sequence.<br> For example:

```
"Bat"
'I am Walking to the Store'
'Data Science is #1'
```

### Basic Rules of Strings:
- Zero or More Characters
- The Characters are in Single or Double Quotation Marks

**Question 2.2.1**<br> Write your own string by assigning it to the variable `my_string`; we will discuss more about variable assignment in **2.4**

```python
my_string = ...
my_string
```

```python
grader.check("q2.2.1")
```

<a id="int_float"></a>
### 2.2.2 Integers and Floating Point Numbers
Numerical data are either integers or floating point numbers.

- Integers
    - Numbers with no decimal
    - Ex: Count of different species within certain habitat
- Floating Point Number
    - Numbers with a decimal
    - Ex: Experimental data collection related to weights or other empirical forms of measure.
    - Used for greater precision.


```python
# Integer
5
```

```python
# Float
10.0
```

```python
# Float 
.23
```

<a id="exp"></a>
## 2.3 Expressions

Expressions are a combination of values, variables and [arithmetic operators](https://inferentialthinking.com/chapters/03/1/Expressions.html)

Here we can see an expression showcasing the subtraction of two floats.

```python
3.25-1.5
```

Many basic arithmetic operations are built into Python.  The textbook section on [Expressions](http://www.inferentialthinking.com/chapters/03/1/expressions.html) describes all the arithmetic operators used in the course.  The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8. 

The order of operations is the same as in math, and Python also has parentheses. For example, compare the outputs of the cells below. The second cell uses parentheses for the current year.


```python
3+6*5-6*3**2*2**3/4*7
```

```python
8+(6*5-(6*3))**2*((2**3)/4*7)
```

In standard math notation, the first expression is

$$3 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$7 + (6 \times 5 - (6 \times 3))^2 \times (\frac{(2^3)}{4} \times 7).$$

**Question 2.3**<br> Write a Python expression in this next cell that's equal to $5 \times (3 \frac{10}{11}) - 50 \frac{1}{3} + 2^{.5 \times 22} - \frac{7}{33} + 7$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus seven. By "$3 \frac{10}{11}$" we mean $3+\frac{10}{11}$, not $3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary. Write this expression after the `=` this will assign the expression to the variable year. This will be explained further in the following section.

*Hint:* The correct output should start with a familiar number.

```python
year = ...
year
```

```python
grader.check("q2.3")
```

<a id="var"></a>
## 2.4 Variable Assignment 

In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

**Question 2.4**<br> Try writing code that uses a name (like `eleven`) that hasn't been assigned to anything.  You'll see an error!

```python

```

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.


```python
a = 841.25
b = 2 * a
c = 12
d = c * b
d
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.


Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
semi_month_sal = 841.25 # $, Semimonthly salary
month_sal = 2 * semi_month_sal # Monthly Salary
num_month_yr = 12 # Number of Months Per Year
year_sal = num_month_yr * month_sal # Yearly Salary
year_sal
```

**However**, names are very important for making your code *readable* to yourself and others.  The first cell with variable a,b,c,d has shorter variables, but they are totally useless without an explanation of what each of them truly mean. 

In the above cell we can see variable names that are relatively short but also descriptive of their purpose.

Also including comments wherever possible can aid in the description of your variables to keep them short in length.

<a id="q3"></a>
# 3. Question 3<br>
Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.



```python
# Change the next line 
# so that it computes the number of seconds in a decade 
# and assigns that number the name, seconds_in_a_decade.

seconds_in_a_decade = ...

# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("q3")
```

<a id="q4"></a>
# 4. Question 4
Define two strings that are four letter English words that include *only* the letters `a,g,n,p,r,s,t`. You can repeat a letter and use plurals.



```python
word_one= ...
word_two= ...
```

```python

```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.6 (points: 2)
```
>>> assert isinstance(test_answer, int)
```
failure_message: Answer should be a integer, not a string or float
```
>>> assert test_answer == 4
```
failure_message: Answer of 2+2 should be 4


Tests for q2.2.1 (points: 2)
```
>>> assert isinstance(my_string, str)
```
failure_message: Answer should be a string


Tests for q2.3 (points: 2)
```
>>> assert 'year' in vars()
```
failure_message: Variable `year` doesn't exist, re-run cell or ensure variable hasn't been deleted
```
>>> assert year != ...
```
failure_message: No expression or value has been set to `year`
```
>>> assert year == 2024
```
failure_message: Answer should be 2023


Tests for q3 (points: 2)
```
>>> assert 'seconds_in_a_decade' in vars()
```
failure_message: Variable doesn't exist, re-run cell, or ensure variable hasn't been deleted
```
>>> assert seconds_in_a_decade != ...
```
failure_message: Make sure to assign expression to variable
```
>>> assert seconds_in_a_decade != 315360000
```
failure_message: 2 leap years and 8 non-leap years
```
>>> assert seconds_in_a_decade == 315532800
```
failure_message: answer should be:315532800


Tests for q4 (points: 2)
```
>>> assert isinstance(word_one, str)
>>> assert isinstance(word_two, str)
```
failure_message: Are you setting word_one and word_two to be strings?
```
>>> assert len(word_one) == 4
>>> assert len(word_two) == 4
```
failure_message: Words need to have a length of 4 characters
```
>>> letter_set = ['a', 'g', 'n', 'p', 'r', 's', 't']
>>> def contains_all_test(test_word, letter_set):
...     for cr_letter in test_word:
...         if cr_letter not in letter_set:
...             return False
...         return True
>>> assert contains_all_test(word_one, letter_set)
>>> assert contains_all_test(word_two, letter_set)
```
